Перед сдачей обязательно убедитесь, что всё работает как надо. Во-первых, **перезапустите ядро и запустите все ячейки** (меню Kernel$\rightarrow$Restart & Run All).

Укажите ваше имя и имена тех, с кем вместе вы решали это ДЗ. Проверьте, что во всех ячейках есть ответы.

```python
NAME = ""
COLLABORATORS = ""
```

---

# Домашнее задание 2
Дедлайн 29.10.2019 23:59

## Задание 1(1 балл)

Напишите функцию, которая принимает список, и возвращает новый список, состоящий из уникальных элементов данного списка. Порядок элементов должен сохраниться

__________________________________________

**Пример входных данных:**  
list_no_dubs([8, 'a', 154, 8, 12])

**Пример выходных данных:**  
[8, 'a', 154, 12]

<!-- grade_id: cell-25d318e1303cb6d8, points: 0, grade: false, solution: true -->
```python
def list_no_dubs(lst):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-3582972d2a678fab, points: 1, grade: true, solution: false -->
```python
assert list_no_dubs([8, 'a', 154, 8, 12]) == [8, 'a', 154, 12]
```

## Задание 2 (2 баллa)

Напишите функцию, которая переводит римские цифры в арабские. 

Напомним, как работать с римскими цифрами. Каждая буква обозначает свое число
    
* I: 1
* V: 5
* X: 10
* L: 50
* C: 100
* D: 500
* M: 1000

Кроме этого, есть **особый случай**, когда цифра меньшего порядка идет до цифры большего порядка. Это означает, что данное число вычитается из предыдущего. Например, число 8 можно записать как VIII, а число 9 уже записывается как IX, а **не** XIIII.

Арабские буквы могут быть только заглавные и только из списка выше, проверять это не нужно.

Пустая строка должна возвращать 0.

__________________________________________

**Пример входных данных:**  
to_arabic("IIX")

**Пример выходных данных:**  
8

<!-- grade_id: cell-c94e27fccd166aa7, points: 0, grade: false, solution: true -->
```python
def to_arabic(number):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-173b15310162d38a, points: 2, grade: true, solution: false -->
```python
assert to_arabic('I') == 1
assert to_arabic('III') == 3
```

## Задание 3 (2 баллa)

Напишите функцию, которая переводит арабские цифры в римские. Подаваемое число не может быть больше 3999.

*HINT*: в [русской википедии](https://ru.wikipedia.org/wiki/%D0%A0%D0%B8%D0%BC%D1%81%D0%BA%D0%B8%D0%B5_%D1%86%D0%B8%D1%84%D1%80%D1%8B) вы можете найти полезную подсказку про то, какие бывают случаи "вычитания" в римских цифрах.


__________________________________________

**Пример входных данных:**  
to_roman(8)

**Пример выходных данных:**  
IIX

<!-- grade_id: cell-3bde515a28696809, points: 0, grade: true, solution: true -->
```python
def to_roman(number):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-41959e58df00f84c, points: 2, grade: true, solution: false -->
```python
assert to_roman(1) == 'I'
assert to_roman(3) == 'III'
```

## Задание 4 (1 балл)

Напишите программу, которая принимает 5 чисел, и возвращает медиану этих чисел: два числа должны быть больше данного, а два меньше.
   
*Дополнительный балл:* используйте менее 7 сравнений.
    

__________________________________________

**Пример входных данных:**  
median(1,4,2,5,3)

**Пример выходных данных:**  
3

<!-- grade_id: cell-0cc92accc9cf8d95, points: 0, grade: false, solution: true -->
```python
def median(a,b,c,d,e):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-a117855cb6bb3737, points: 1, grade: true, solution: false -->
```python
assert median(1,7,546,5,2) == 5
```

## Задание 5 (2 балла)

Вы находитесь на шоу "Давай договоримся!". Ведущий показывает вам три закрытые двери. За одной из них шикарный приз, а за двумя другими пусто. Ведущий знает где находится приз, а вы нет. Вы выбираете одну из дверей, но не открываете её! Далее ведущий, знающий, где находится приз, открывает одну из дверей без приза, и предлагает вам выбор: либо **изменить** свой выбор, либо **остаться при своем**.

Какая из двух стратегий ("поменять дверь" и "остаться при своем") будет иметь большую вероятность победы?

Напишите программу, которая симулирует обе стратегии. Через `input` программа должна принимать число `n` -- количество симуляций. Программа должна `n` раз сыграть первую стратегию, и `n` раз сыграть вторую стратегию.

В задаче вам может понадобиться сгенерить случайное число. Это можно сделать при помощи функции `random.randint`.

```
from random import randint

randint(0, 2)
# возвращает 0, 1, 2 равновероятно
```

<!-- grade_id: cell-ed004b8359390ff8, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python

```

## Задание 6 (2 балла)

Напишите код, иллюстрирующий простую модель роста популяции, которую можно встретить в исследованиях рыб в водоеме, бактерий в чашке Петри. Предположим популяция может быть между 0 (полное вымирание) и 1 (макксимальная популяция). Если в момент времени `t` популяция равна $x$, то в момент времени `t+1` популяция будет равна $r \cdot x \cdot (1-x)$, где $r$ это *параметр плодовитости*, обозначающий скорость роста. Начнем с небольшой популяции в $x = 0.01$, исследуйте модель для разных значений $r$. Для каких значений $r$ размер популяции стабилизируется в точке $x = 1-1/r$? Что вы можете сказать о популяции, когда $r$ равно 3.5? 3.8? 5?

<!-- grade_id: cell-cc46ab7499ff1917, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python

```

```python

```

```python

```